Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Alekseev Arsalan"
COLLABORATORS = ""
```

---

```python
import sys

import numpy as np
import unittest
import time

import collections
import pickle
import io
```

<!-- grade_id: cell-ac8fc52d8a39ccb4, points: 0, grade: false, solution: true -->
```python
from k_nearest_neighbor import KNearestNeighbor

#class KNearestNeighbor:
    # YOUR CODE HERE
    #raise NotImplementedError()
```

### Test 1: Two loops (0.05 points)

<!-- grade_id: cell-3473b7b6ffd64d07, points: 0.01, grade: true, solution: false -->
```python
# do not change this cell
knn_test = KNearestNeighbor()


```

<!-- grade_id: cell-e3503c286039ec55, points: 0.05, grade: true, solution: false -->
```python
def time_function(f, *args):
    """
    Call a function f with args and return the time (in seconds) that it took to execute.
    """
    tic = time.time()
    output = f(*args)
    toc = time.time()
    return toc - tic, output

knn_test.fit(X_train, y_train)

np.random.seed(42)
two_loop_time, out_2_loops = time_function(knn_test.compute_distances_two_loops, X_test)
assert np.allclose(real_distances, out_2_loops, atol=1e-6)
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[6], line 10
      7     toc = time.time()
      8     return toc - tic, output
---> 10 knn_test.fit(X_train, y_train)
     12 np.random.seed(42)
     13 two_loop_time, out_2_loops = time_function(knn_test.compute_distances_two_loops, X_test)

NameError: name 'X_train' is not defined
```

### Test 2: One loop (0.1 points)

<!-- grade_id: cell-e2c4124a6f815118, points: 0.1, grade: true, solution: false -->
```python
np.random.seed(42)
one_loop_time, out_1_loops = time_function(knn_test.compute_distances_one_loop, X_test)
assert np.allclose(real_distances, out_1_loops, atol=1e-6)
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[22], line 2
      1 np.random.seed(42)
----> 2 one_loop_time, out_1_loops = time_function(knn_test.compute_distances_one_loop, X_test)
      3 assert np.allclose(real_distances, out_1_loops, atol=1e-6)

NameError: name 'X_test' is not defined
```

### Test 3: No loops (0.4 points)

<!-- grade_id: cell-69473387a23d8dff, points: 0.4, grade: true, solution: false -->
```python
np.random.seed(42)
no_loop_time, out_no_loops = time_function(knn_test.compute_distances_no_loops, X_test)
assert np.allclose(real_distances, out_no_loops, atol=1e-6)

```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[23], line 2
      1 np.random.seed(42)
----> 2 no_loop_time, out_no_loops = time_function(knn_test.compute_distances_no_loops, X_test)
      3 assert np.allclose(real_distances, out_no_loops, atol=1e-6)

NameError: name 'X_test' is not defined
```

### Test 4: No loops timing (0.15 points)

<!-- grade_id: cell-3460671ba79fd04a, points: 0.15, grade: true, solution: false -->
```python
np.random.seed(42)
two_loop_time, out_2_loops = time_function(knn_test.compute_distances_two_loops, X_test)
no_loop_time, out_no_loops = time_function(knn_test.compute_distances_no_loops, X_test)
assert np.log(two_loop_time) - np.log(no_loop_time) > np.log(50)
```

### Test 5: Labels prediction (0.3 points)

<!-- grade_id: cell-5d56bb4222ee1e7c, points: 0.3, grade: true, solution: false -->
```python
np.random.seed(42)
for k in [1, 3, 5, 7, 9]:
    predicted_labels = knn_test.predict(X_test, k=k)
    predicted_labels = np.array(predicted_labels, dtype=int).squeeze()
    real_labels = np.array(y_ref_predictions[k], dtype=int).squeeze()
    assert np.array_equal(predicted_labels, real_labels), 'Wrong answer for k={}'.format(k)
```